```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_2_functions.ipynb")
```

# Lecture goals

1. (more) Dictionaries for data encapsulation
2. Functions for functionality encapsulation

Lab slides: https://docs.google.com/presentation/d/1ykwwcQ0onMvAjUxfJmKl9tbo-rJPdB5pRwDEmpDsd-g/edit?usp=sharing

## Functions: 
Functions enable encapsulation of, well, functionality.

They're also a useful mental tool for organizing and structuring your thoughts on how to solve a given problem
1. Clearly define a bit of code that takes in some inputs, does some computation, then outputs some data
2. Makes it easier to test that code with different inputs
3. Practicalities: Prevents one of the most common sources of errors - re-using variable names

It's almost never wrong to encapsulate a bit of code in a function. It can slow down (a tiny bit) computation time, but can greatly reduce debugging time, so it's usually worth it.

Python's function syntax is beautifully designed to make it easy to set default values for parameters and pass back as much data as you want. We'll see more of that later; for this assignment we'll use the power of dictionaries to pass back "labeled" data.

In this lecture activity you're essentially going to copy over the code you did in the previous lecture activity, right-shift it to place it in the function, replace the **test_** variable name with the input parameter to the function, then add a return to return the answer. This is a pretty common approach for turning code into a function. 


```python
# Access all numpy functions as np.
import numpy as np
```

## Question 1: Stats on a list

### Calculate stats on a list

TODO: in calc_stats_from_list function 
- Calculate the mean of the negative and positive values
- Count the total number of negative/positive values
- Store the values in a dictionary

This function calculates the given stats from the list that is passed in. There is test code below this function.

TODO: 
 - Step 1 - copy your code from lecture activity 1 into the function. Right shift it so it is indented properly
 - Step 2 - change **test_list_one** to be **in_list** - the input parameter of the function
 - Step 3 - return the dictionary **dict_save_stats**

```python
def calc_stats_from_list(in_list):
    """ Calculate mean of positive numbers, mean of negatives numbers
    Separate the list into positive and negative numbers. Calculate the mean of each. Return those means, along with
     how many positive/negative numbers there were
    @param in_list : any list type
    @return - A dictionary with the desired stats"""

    # These are the stats we're calculating. This is more elegant/useful than creating four variables - it keeps all
    #  of the values in the same place and assigns a meaningful label (key) to them
    dict_save_stats = {"Mean positive": 0.0, "Mean negative": -0.0, "Count positive": 0, "Count negative": 0}

    # TODO: 
    #   Copy your code from lecture activity 1 here. Don't forget to change the name of test_list_one to be in_list
    for item in in_list: #iterates over each item in test_list_one, I used example code to start
        if item < 0: #if the item is negative do the below steps
            dict_save_stats["Count negative"] += 1 #addition assignment calculator adds to count negative
            dict_save_stats["Mean negative"] += item #adds an item to the mean calculation
        elif item > 0: #if the item is positive do the below steps
            dict_save_stats["Count positive"] +=1 #same as when negative but going towards the positive steps
            dict_save_stats["Mean positive"] += item
        
    if dict_save_stats["Count negative"]: #collects the negative counts
        dict_save_stats["Mean negative"] /= dict_save_stats["Count negative"] #this divides the sum of the negative mean and divides it by the negative count
    if dict_save_stats["Count positive"]: #collects the positive counts
        dict_save_stats["Mean positive"] /= dict_save_stats["Count positive"] #same step as the neagtive mean
    
    # TODO Do the return here
    #return function returns result 
    return dict_save_stats

```

### Test code for list

Create the arrays and test your function using those test arrays. Here's another advantage of functions - you can create test data for yourself to make sure the code is working right. Encapsulating the code in a function means you don't 
1. Accidentally change the code when switching from the test data to the real data
2. You can make more than one test 
3. You can run the tests more than once/all the time to double check that you didn't "break" the code

TODO: 
- Fill in the calc_stats_from_list function above
- Run the cell below - it will print out if your values are incorrect

Note that, below, we'll test this code one last time with randomly generated data

```python
# Test data
test_list_one = [-0.75, -0.25, 1.0 / 3.0, 2.0 / 3.0, 3.0 / 3.0]
test_list_res = calc_stats_from_list(test_list_one)

b_tests_passed = True
if not np.isclose(test_list_res["Mean positive"], 2.0 / 3.0):
    b_tests_passed = False
    print(f"Mean positive is not correct, should be {2.0/3.0}, got {test_list_res['Mean positive']}")

if not np.isclose(test_list_res["Mean negative"], -0.5):
    b_tests_passed = False
    print(f"Mean negative is not correct, should be -0.5, got {test_list_res['Mean negative']}")

if test_list_res["Count positive"] != 3:
    b_tests_passed = False
    print(f"Count positive numbers, should be 3, got {test_list_res['Count positive']}")

if test_list_res["Count negative"] != 2:
    b_tests_passed = False
    print(f"Count positive numbers, should be 2, got {test_list_res['Count negative']}")

if b_tests_passed:
    print("All array tests passed!")
```

Output:
```
All array tests passed!
```

```python
# This is an example of what the autograder tests are doing
res = calc_stats_from_list([-1, 2, -3, 4, -5])
assert np.isclose(res["Mean positive"], 3.0)

# Remember that if you are printing anything out in calc_stats_from_list then the autograder will fail
```

```python
grader.check("list")
```

Output:
```
list results: All test cases passed!
```

# Question 2: Doing it again with a numpy array

## Fill in calc_stats_from_nparray

For this function, assume the input is a numpy array.

TODO: Same as the previous question, but this time copy in the numpy array code you wrote in lecture activity 1. Again:
- NO **if** statements or **for** loops - do this all with numpy operations

As before, test code is below

```python
def calc_stats_from_nparray(in_nparray):
    """ Calculate mean of positive numbers, mean of negatives numbers
    Separate the list into positive and negative numbers. Calculate the mean of each. Return those means, along with
     how many positive/negative numbers there were
    @param in_nparray : numpy array
    @return - A dictionary with the desired stats"""

    # TODO: Copy in the code, change the numpy array name to be the input to this function, and return the dictionary
    dict_save_stats_np = {"Mean positive": 0, "Mean negative": 0, "Count positive": 0, "Count negative": 0}
#count_nonzero allows you to count nonzero array's values
    Count_positive=np.count_nonzero(in_nparray>0) # number of psoitive values
    Count_negative=np.count_nonzero(in_nparray<0) # number of negative values

#have to use test_nparry> or< 0 not just count positive (mistake I made first)
#calculating mean and np.mean allows you to do this
    Mean_positive= np.mean(in_nparray[in_nparray>0]) #np.mean calculated mean and stuff inside () filters out the values that are positive
    Mean_negative= np.mean(in_nparray[in_nparray<0]) #same steps as above but negative

#dictionary needs to be updated (forget this step to get grader check right)
#confused when the mean_positive switches places with the other side because my code didnt work when I had this
    dict_save_stats_np["Mean positive"]=Mean_positive
    dict_save_stats_np["Mean negative"]=Mean_negative
    dict_save_stats_np["Count positive"]=Count_positive
    dict_save_stats_np["Count negative"]=Count_negative

#copied from question 1
#return function returns result 
    return dict_save_stats_np
```

### Test code for numpy array

This will print out if your function above is returning incorrect values

```python
test_nparray_one = np.array(test_list_one)  # Convert the previous test list to a numpy array
test_list_res = calc_stats_from_nparray(test_nparray_one)

b_tests_passed = True
if not np.isclose(test_list_res["Mean positive"], 2.0 / 3.0):
    b_tests_passed = False
    print(f"Mean positive is not correct, should be {2.0/3.0}, got {test_list_res['Mean positive']}")

if not np.isclose(test_list_res["Mean negative"], -0.5):
    b_tests_passed = False
    print(f"Mean negative is not correct, should be -0.5, got {test_list_res['Mean negative']}")

if test_list_res["Count positive"] != 3:
    b_tests_passed = False
    print(f"Count positive numbers, should be 3, got {test_list_res['Count positive']}")

if test_list_res["Count negative"] != 2:
    b_tests_passed = False
    print(f"Count positive numbers, should be 2, got {test_list_res['Count negative']}")

if b_tests_passed:
    print("All numpy array tests passed!")
```

Output:
```
All numpy array tests passed!
```

```python
grader.check("nparray")
```

Output:
```
nparray results: All test cases passed!
```

### Encapsulating the data slice 

In this problem we'll encapsulate the fancy slice you did in lab 2 where you get out the x,y,z columns for the wrist torque data. 

```python
# Load in the data
import json as json
try:
    fname = "Data/proxy_pick_data.csv"
    pick_data = np.loadtxt(fname, dtype="float", delimiter=",")
except FileNotFoundError:
    print(f"File not found {fname}")


# TODO - copy over your code from Lab 1 here
try:
    with open("Data/proxy_data_description.json", "r") as fp:
        pick_data_description = json.load(fp)
        data_channels=pick_data_description["Data channels"]
        print(f"Data channels:{data_channels}")
except FileNotFoundError:
    print(f"The file was not found; check that the data directory is in the current one and the file is in that directory")
n_picks = np.count_nonzero(pick_data[:, 1]) 
n_successful = np.count_nonzero(pick_data[:, -1] == 1)
n_dimensions=pick_data.shape[0]
n_time_steps=pick_data.shape[1]
n_total_dims = 0
data_channels=pick_data_description["Data channels"]

for item in data_channels:
    n_dimensions = item ["dimensions"] #get number of dimension for channel
    n_total_dims += n_dimensions #add number of dimensions
pick_channel_data = pick_data[:,:-1] 
pick_successful = pick_data[:,-1] 
```

Output:
```
Data channels:[{'name': 'Wrist force', 'index_offset': 0, 'dimensions': 3, 'units': 'N'}, {'name': 'Wrist torque', 'index_offset': 3, 'dimensions': 3, 'units': 'N.m'}, {'name': 'IMU accel f1', 'index_offset': 6, 'dimensions': 3, 'units': 'gravity units'}, {'name': 'IMU accel f2', 'index_offset': 15, 'dimensions': 3, 'units': 'gravity units'}, {'name': 'IMU accel f3', 'index_offset': 24, 'dimensions': 3, 'units': 'gravity units'}, {'name': 'IMU velocity f1', 'index_offset': 9, 'dimensions': 3, 'units': 'deg/sec'}, {'name': 'IMU velocity f2', 'index_offset': 18, 'dimensions': 3, 'units': 'deg/sec'}, {'name': 'IMU velocity f3', 'index_offset': 27, 'dimensions': 3, 'units': 'deg/sec'}, {'name': 'Motor position f1', 'index_offset': 12, 'dimensions': 1, 'units': 'radians'}, {'name': 'Motor position f2', 'index_offset': 21, 'dimensions': 1, 'units': 'radians'}, {'name': 'Motor position f3', 'index_offset': 30, 'dimensions': 1, 'units': 'radians'}, {'name': 'Motor velocity f1', 'index_offset':
```

```python
def get_channel_data(all_data, n_picks, start_index, n_time_steps, n_total_dims, n_dims):
    """ Get the data for just one channel (eg, wrist torque)
    @param all_data - the pick_channel_data numpy array
    @param n_picks - number of picks (number of rows in all_data)
    @param start_index - where to start getting data from 
    @param n_time_steps - number of time steps
    @param n_total_dims - what the skip value is - the total number of channels
    @param n_dims - total number of dimensions to use (1, 2, or 3)
    @return Return array should be n_picks X (n_timesteps * n_dims)"""

    # TODO Your slice code goes here. Note that I kept most of the variable names the same, so you should only have
    #  to change the wrist torque specific ones

    wrist_torque_data = np.zeros((int(n_picks), int(n_time_steps * n_dims)))

    for item in range(0,3):
        wrist_torque_data[:, item::start_index] = all_data[:, start_index + item::n_total_dims]
    return (wrist_torque_data)
    
```

```python
# Just for this problem I'm "hard-wiring" in these values - this is because I want to test the function
# with known values
n_picks = 660
wt_start_index = 3  # Wrist torque starts at 3
n_time_steps = 40
n_total_dims = 33

# Note - I know the wrist torque data has 3 dimensions (x,y,z)
wrist_torque_data = get_channel_data(pick_channel_data, 
                                    n_picks=n_picks, 
                                    start_index=wt_start_index,
                                    n_time_steps=n_time_steps,
                                    n_total_dims=n_total_dims,
                                    n_dims=3)
```

```python
# SELF TESTS
# If you need to, copy over the tests from lab 1, the slicing problem - the values should be the same
```

```python
grader.check("data_function")
```

Output:
```
data_function results: All test cases passed!
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"Kaleb"}
# List of URLS TCS3 (creates a set)
websites = {"na"}
# Approximate number of hours, including lab/in-class time
hours = 3
```

```python
grader.check("hours_collaborators")
```

Output:
```
hours_collaborators results: All test cases passed!
```

### To submit

- Do a restart then run all to make sure everything runs ok
- Save the file
- Submit just this .ipynb file through gradescope, Lecture activity 2, functions
- You do NOT need to submit the data files - we will supply those

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Most likely failure for this assignment is not naming the data directory and files correctly; capitalization matters for the Gradescope grader. 

Tests for data_function (points: 2)
```
>>> assert(wrist_torque_data.shape[0] == pick_data.shape[0])
```
```
>>> assert(wrist_torque_data.shape[1] == 3 * 40)
```
```
>>> assert(np.isclose(wrist_torque_data[0, 0], -0.268183058, atol=0.001))
```
```
>>> assert(np.isclose(wrist_torque_data[0, -1], -0.-0.077284199, atol=0.001))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for list (points: 1)
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Mean positive"], 3.0)
```
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Mean negative"], -3.0)
```
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Count positive"], 2)
```
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Count negative"], 3)
```


Tests for nparray (points: 1)
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Mean positive"], 3.0)
```
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Mean negative"], -3.0)
```
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Count positive"], 2)
```
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Count negative"], 3)
```
